# BLU 2  - Exercise Notebook 

```python
import pandas as pd
import matplotlib.pyplot as plt
plt.rcParams['figure.figsize'] = (16, 8)
import warnings
warnings.simplefilter(action='ignore', category=FutureWarning)   
from random import seed

from sklearn.linear_model import LinearRegression
import numpy as np
import math
from statsmodels.tsa.seasonal import seasonal_decompose
from statsmodels.graphics.tsaplots import plot_acf, plot_pacf
from statsmodels.tsa import stattools
from pandas.plotting import lag_plot
from statsmodels.tsa.arima_model import ARIMA

% matplotlib inline
```

# Airlines! 

You will be working on the classic Airlines dataset. We'll do the pre-processing for you, so that you can focus on the predictions part. 

```python
def load_airline_data():
    airlines = pd.read_csv('Learning Notebooks/data/international-airline-passengers.csv', 
                       index_col='Month')[:-1]

    airlines.columns = ['passengers_thousands']
    airlines = airlines['passengers_thousands']
    airlines.index = pd.to_datetime(airlines.index)
    
    return airlines
```

```python
airlines = load_airline_data()
```

```python
airlines.head()
```

##### Plot the series to get an idea of what's going on

```python
airlines.plot(figsize=(16, 6));
```

### Q0 (ungraded): 
 - Is the magnitude of the variance changing? 
 - Does it look cyclical? 
 - Does it have an apparent trend? 

## Let's identify our time series components

### Q1: Estimate the trend of the time series using Linear Regression:

<!-- grade_id: cell-7b2d7a5eee8d277f, points: 0, grade: false, solution: true -->
```python
def plot_linear_regression(df_):
    
    # Hint: we did this in BLU2 Part 2 
    
    # X = # create an array with the steps 
    # lr = # you Linear Regression 
    # fit your logit  
    # linear_trend = # predict the data 
    
    # YOUR CODE HERE
    raise NotImplementedError()

    linear_trend.plot();
    df_.plot();
    plt.show()
    
    return linear_trend    

linear_trend = plot_linear_regression(airlines)
```

<!-- grade_id: cell-8bafcb20ee1977d9, points: 2, grade: true, solution: false -->
```python
assert math.isclose(linear_trend.sum(), 40363.0)
```

### Q2: Estimate the trend of the time series using Rolling window

<!-- grade_id: cell-485e5fa4abd88e13, points: 0, grade: false, solution: true -->
```python
def plot_rolling_windows(df_):
    
    # rolling_6 =     # create a rolling window of 6 months
    # rolling_12 =    # create a rolling window of 12 months

    # YOUR CODE HERE
    raise NotImplementedError()

    df_.plot()
    rolling_12.plot()
    rolling_6.plot()
    plt.show()
    
    return rolling_6, rolling_12

rolling_6, rolling_12 = plot_rolling_windows(airlines)
```

<!-- grade_id: cell-0803a24329962f61, points: 2, grade: true, solution: false -->
```python
assert rolling_6.isnull()[0:5].sum()
assert math.isclose(rolling_6.sum(), 38941.0)
assert math.isclose(rolling_12.mean(), 278.17731829)
```

### Q3: Check the variance 

<!-- grade_id: cell-dbc5a65e6dd90854, points: 0, grade: false, solution: true -->
```python
def observing_variance(df_):
    
    
    # variance_verification = # use a rolling window of 24 months to see what's happening with the standar deviation
    # Then plot the variance_verification 

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return variance_verification

variance_verification = observing_variance(airlines)
```

<!-- grade_id: cell-291c03ac60c68f7d, points: 2, grade: true, solution: false -->
```python
assert math.isclose(variance_verification.sum(), 5302.061403933007)
```

### Q4: Make a seasonal decomposition

##### Remember we have shown you two models: 'additive' and 'multiplicative' . Which one should you use? 

<!-- grade_id: cell-c369f965ee868273, points: 0, grade: false, solution: true -->
```python
def decomposition(df_): 
    
    # uncomment the model you wish to use
    # model = 'additive'
    # model = 'multiplicative' 

    # YOUR CODE HERE
    raise NotImplementedError()


    # decomp = # use statsmodel to get a seasonal decomposition object 
    # YOUR CODE HERE
    raise NotImplementedError()

    # plot the observed, trend, seasonal and residuals from decomp
    # YOUR CODE HERE
    raise NotImplementedError()

    plt.show()
    
    return decomp 

decomp = decomposition(airlines)
```

<!-- grade_id: cell-567790cdbf31324b, points: 2, grade: true, solution: false -->
```python
assert math.isclose(decomp.trend.sum(), 36696.166666666664)
assert math.isclose(decomp.seasonal.sum(), 144)
```

### Q5: Create a stationary series
- What is the detrended series? (Use a lag one for the difference)
- Plot the detrended series
- What's the maximum correlation (except the first day) 
- What interpretations can be done from here? 

<!-- grade_id: cell-c14a426948b6b076, points: 0, grade: false, solution: true -->
```python
def make_stationary(df_):
    # Here are some steps to help you:
    # log_df = # remove the increase of the magnitude in the variance 
    # detrend = # get the difference (use the log_df!)
    # detrend = # drop the nulls to avoid hell
    # plot detrend 
    # Think about what you are seeing! The mean and the variance of the series you are looking at

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return log_df, detrend

log_df, detrend = make_stationary(airlines)
```

<!-- grade_id: cell-e8322815bf956221, points: 2, grade: true, solution: false -->
```python
assert math.isclose(detrend.sum(), 1.3499267169490166)
assert np.exp(log_df).values.all() == airlines.values.all()
```

This looks constant in time: Variance and mean! This is going towards stationary, hurray! 

Finally... we have some rest from those trends and variations now it's time to start working with the simple things of this learning unit! 

### Q6: Get me some autocorrelations please! 
1. Get the autocorrelation and partial autocorrelation for the series

<!-- grade_id: cell-b2bfb6431053bbb1, points: 0, grade: false, solution: true -->
```python
def get_acf_and_pacf(df_):

    # acf_result = # series of the autocorrelation 
    # pacf_result = # series of the partial autocorrelation 
    
    # YOUR CODE HERE
    raise NotImplementedError()

    # plot the autocorrelation (use statsmodel)
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # plot the partial autocorrelation (use statsmodel)
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return acf_result, pacf_result


# Call the function with the right argument. 
# Is the right argument (df_) airplanes? Probably not, or we wouldn't be asking. 

# acf_result, pacf_result = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a732134f5ff90a6c, points: 2, grade: true, solution: false -->
```python
assert math.isclose(acf_result.sum(), 0.6453106068705969)
assert math.isclose(pacf_result.sum(), -1.901056081553599)
```

### Q7 interpret the plots 

<!-- grade_id: cell-c2fe0fd6e61d852e, points: 0, grade: false, solution: true -->
```python
# relevant_lags_acf = # the number of lags, except zero, 
#     that are statistically significant before crossing the first time the blue bondary

# relevant_lags_pacf = # the number of lags, except zero, 
#     that are statistically significant before crossing the first time the blue bondary

# you can hardcode this one: 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1d7bdfbd336470e2, points: 2, grade: true, solution: false -->
```python
# we are giving this one easy, as it is more important that you know you got it right. 
# however, don't just copy the numbers, if you don't understand where they come from ask us! 
assert relevant_lags_acf == 1
assert relevant_lags_pacf == 2
```

### Q8: Get the lag scatter for the lag 2 and 12. Which one seems to show seasonality? 

Hint: this is also in Part 2 of the Learning Notebooks 

<!-- grade_id: cell-9d52ada83231f246, points: 0, grade: false, solution: true -->
```python
# Plot the lag scatter for lag 2 

# YOUR CODE HERE
raise NotImplementedError()
plt.show()

# Plot the lag scatter for lag 12 
# YOUR CODE HERE
raise NotImplementedError()
plt.show()

# uncomment the one you think is true (from observing the plots) 
# lag = 12
# lag = 2

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-dc3b0fb4842ec423, points: 2, grade: true, solution: false -->
```python
assert lag == 12
```

### Q9: Choose the parameters for your ARIMA model based on what you have seen before on the autocorrelation

Note 1: you will need to understand what the d is doing and use previous answers to select `p`, `d` and `q`

_Note: In case of doubt for the **lag number 2** in the Partial Autocorrelation Function, consider that it crosses the blue/dashed lines_

<!-- grade_id: cell-6312e30d5d2437ff, points: 0, grade: false, solution: true -->
```python
def fit_arima(df_, p, d, q): 
    # model = ... 
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # Actually fit the model (remember, statsmodel is different from scikit in API)
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return results


# Chose the parameters, and train the model 
# p = ... 
# d = ...
# q = ...
# results = ...

# YOUR CODE HERE
raise NotImplementedError()

# plot both series together 
# plot the fitted values
# plot the series you are comparing (again, not telling you which one), 
#     but remember arima works with stationary series (wink, wink) 


# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cbff638ed2886e52, points: 2, grade: true, solution: false -->
```python
assert math.isclose(results.fittedvalues.sum(), 0.612070865961366, rel_tol=0.001)
```

### Q10: Get a plot with the original series and the one that is being given by the ARIMA you have built

This really well behaved series are cool but we don't please our bosses with this. Let's get our original series and see how our ARIMA is behaving

<!-- grade_id: cell-da752e940385af51, points: 0, grade: false, solution: true -->
```python
def get_original_series_back(model_results): 
    
    # pred_arima = # get a series with the values the arima is giving
    # pred_arima_cumsum = # get a cummulative sum of the differences

    # pred_arma_log = # get a series with just the step 0 of the logs (the step before diff)
    # pred_arma_log = # add to the logs the cum sum 

    # hello_old_friend = # just need to revert the log part
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return pred_arima, pred_arima_cumsum, pred_arma_log, hello_old_friend


pred_arima, pred_arima_cumsum, pred_arma_log, hello_old_friend = get_original_series_back(results)

# uncomment to see both series together

# plt.plot(airlines.sort_index(), label='original')
# plt.plot(hello_old_friend.sort_index(), label='arima')
# plt.show();
```

<!-- grade_id: cell-2a396328cec2d03d, points: 2, grade: true, solution: false -->
```python
assert math.isclose(pred_arima.sum(), 0.612070865961366, rel_tol=0.01)
assert math.isclose(pred_arima_cumsum.sum(), 11.730354194686626, rel_tol=0.01)
assert math.isclose(pred_arma_log.sum(), 691.1941916611802, rel_tol=0.01)
assert math.isclose(hello_old_friend.sum(), 17990.426122789817, rel_tol=0.01)
```